# Part III: Explore

![alt text](data/inu_neko_logo_small.png "Inu + Neko")

Hello! We are Inu + Neku and we are a Dog & Cat services and supplies store located in New York City. We just started our e-commerce business and need your help analyzing our data!

## Description

We need to make sure the data is clean before starting your analysis. As a reminder, we should check for:

- Duplicate records
- Consistent formatting
- Missing values
- Obviously wrong values

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
```

```python
df_cleaned = pd.read_csv('data/inu_neko_orderline_clean.csv')
df_cleaned
```

Output:
```
trans_id      prod_upc  cust_id trans_timestamp  trans_year  \
0      10300097  719638485153  1001019         35:21.4        2021   
1      10300093   73201504044  1001015         33:37.5        2021   
2      10300093  719638485153  1001015         33:37.5        2021   
3      10300093  441530839394  1001015         33:37.5        2021   
4      10300093  733426809698  1001015         33:37.5        2021   
...         ...           ...      ...             ...         ...   
38218  10327860  287663658863  1022098         37:12.8        2021   
38219  10327960  140160459467  1022157         45:09.9        2021   
38220  10328009  425361189561  1022189         57:44.3        2021   
38221  10328089  733426809698  1022236         59:29.8        2021   
38222  10328109  717036112695  1011924         30:52.2        2021   

       trans_month  trans_day  trans_hour  trans_quantity  cust_age  \
0                1          1           1               1        20   
1                1      
```

```python
df_cleaned.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 38223 entries, 0 to 38222
Data columns (total 16 columns):
 #   Column            Non-Null Count  Dtype  
---  ------            --------------  -----  
 0   trans_id          38223 non-null  int64  
 1   prod_upc          38223 non-null  int64  
 2   cust_id           38223 non-null  int64  
 3   trans_timestamp   38223 non-null  object 
 4   trans_year        38223 non-null  int64  
 5   trans_month       38223 non-null  int64  
 6   trans_day         38223 non-null  int64  
 7   trans_hour        38223 non-null  int64  
 8   trans_quantity    38223 non-null  int64  
 9   cust_age          38223 non-null  int64  
 10  cust_state        38223 non-null  object 
 11  prod_price        38223 non-null  float64
 12  prod_title        38223 non-null  object 
 13  prod_category     38223 non-null  object 
 14  prod_animal_type  38223 non-null  object 
 15  total_sales       38223 non-null  float64
dtypes: float64(2), int64(9), object(5)
memor
```

#### Question 2: Alpha and Omega I
What was the month and day of the first sale? Store as a tuple in that order and assign the tuple to the variable `first_date`.

```python
# 1. Make a new dataframe comprised of variables of interest: 'trans_month', 'trans_day'

df_sorted_ascending = df_cleaned[['trans_month', 'trans_day']]
print(df_sorted_ascending) # Print to verify
```

Output:
```
trans_month  trans_day
0                1          1
1                1          1
2                1          1
3                1          1
4                1          1
...            ...        ...
38218            6         30
38219            6         30
38220            6         30
38221            6         30
38222            6         30

[38223 rows x 2 columns]
```

<!-- grade_id: cell-bd2b89a8035a9ab8, points: 0, grade: false, solution: true -->
```python
# 2. Sort the new dataframe first on 'trans_month' and then on 'trans_day'.
# just like we do in spreadsheet, when we add levels

df_sorted_ascending = df_cleaned[['trans_month', 'trans_day']].sort_values(by = ['trans_month','trans_day'] )
print(df_sorted_ascending) # Print to verify
```

Output:
```
trans_month  trans_day
0                1          1
1                1          1
2                1          1
3                1          1
4                1          1
...            ...        ...
38218            6         30
38219            6         30
38220            6         30
38221            6         30
38222            6         30

[38223 rows x 2 columns]
```

```python
# locate the row with least 'trans_month' and 'trans_day' values
# That means locate the value at index position 0.
first_date = df_sorted_ascending[['trans_month','trans_day']].iloc[0]
print(first_date)
type(first_date)
```

Output:
```
trans_month    1
trans_day      1
Name: 0, dtype: int64
pandas.core.series.Series
```

```python
# convert the first_date to a tuple.
first_date = (tuple(first_date))

first_date
```

Output:
```
(1, 1)
```

# Do the same for Q3 to locate the last value

```python
# Q3 - LATEST ORDER
df_sorted_ascending = df_cleaned.sort_values(by = ['trans_month','trans_day'] )
last_date = tuple(df_sorted_ascending[['trans_month','trans_day']].iloc[-1])
print(last_date)

print(type(last_date))
```

Output:
```
(6, 30)
<class 'tuple'>
```

<!-- grade_id: cell-8a2534037c10eb71, points: 4, grade: true, solution: false -->
```python
# Q2 Test Cases
assert len(first_date) == 2
assert str(first_date[0]).isnumeric() 
assert str(first_date[1]).isnumeric() 
```

```python

```